# BLU06 - Exercises Notebook

```python
import hashlib # for grading

import os
import pandas as pd

from sklearn.base import BaseEstimator, TransformerMixin
from sklearn.datasets import load_boston
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error, mean_absolute_error
from sklearn.model_selection import train_test_split
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import MinMaxScaler
```

## 1 Read the Programs data (graded)

In this first exercise, we aim to create a single data frame, combining all programs from all seasons.

With a caveat though: **we want to include seasons after 1900**.

<!-- grade_id: cell-bf4694e22046237e, points: 0, grade: false, solution: true -->
```python
def make_programs():
    files = os.listdir('data/programs/')
    # Create a list with the name of all files containing programs from
    # 1900 inclusive and onwards (just the filename, no complete path.)
    # files_after_1900: List[str] = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Create a list with the name of all .csv files.
    # seasons: List[pd.DataFrame] = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Use pd.concat to create a single dataframe.
    # programs: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Drop the column ProgramID.
    # programs = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Set the index to be the column GUID, according best practices.
    # Feel free to use method chaining if you want.
    # YOUR CODE HERE
    raise NotImplementedError()
    return programs


def read_season(file):
    path = os.path.join('data', 'programs', file)
    return pd.read_csv(path)


programs = make_programs()
```

<!-- grade_id: cell-1aa243b8e5df1a17, points: 2, grade: true, solution: false -->
```python
assert(programs['Season'].min() == '1900-01')

shape = str(programs.shape)
expected_hash = 'd4e990f328604be98d2ce91ec7cead4660b79c41da0fffa174bd83fc43928c77'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(programs.sample(random_state=0))
expected_hash = '2ed2dbfa0da6dd4ea08fd025b5b744a716c8b47a7a23a06b0cff6198c8e69e94'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

Let's preview the `programs` dataframe.

```python
programs.head()
```

## 2 Read the Concerts data (graded)

Read the concerts data.

Although we list all transformations step-by-step for the sake of clarity, we expect you to use method chaining.

<!-- grade_id: cell-58bd5fde4e09e377, points: 0, grade: false, solution: true -->
```python
def make_concerts(): 
    # Read concerts data and drop the ProgramID and ConcertID columns.
    # concerts: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Remember to_datetime? We need here. We need to parse the columns Date and 
    # Time. Use pd.to_datetime(...).dt.date for the Date and pd_to_datetime(..., 
    # format=%I:%M%p).dt.time for the Time.
    # YOUR CODE HERE
    raise NotImplementedError()
    return concerts


concerts = make_concerts()
```

<!-- grade_id: cell-1e95492b37c4c43f, points: 2, grade: true, solution: false -->
```python
shape = str(concerts.shape)
expected_hash = 'c030586e7370b1f2c34307d5de9b921d96efa28c933e44111b121ed819f339da'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(concerts.sample(random_state=0))
expected_hash = '392a3db01753b02d85173c38cde95112fb5cdf06ca5a45d25f828238d56103be'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

```python
concerts.head()
```

## 3 Combine Programs and Concerts data (graded)

Let's combine both data frames into a single dataset, using an inner join.

<!-- grade_id: cell-110744d4ad3ef2ef, points: 0, grade: false, solution: true -->
```python
# Remember that you want to join on the index of one of the dataframes.
# nyp = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b7730d3cb9832724, points: 2, grade: true, solution: false -->
```python
shape = str(nyp.shape)
expected_hash = 'ed5f695b3d2a0e790163dc73f7fe7e1dc56cc53c6f13289a198702ed4094071c'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(nyp.sample(random_state=0))
expected_hash = '3fa1f19dfd4c9d82b76c963890951b75e1c56a464bc0e592f5f646068ad382d4'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

## 4 Read Works and Soloists data (graded)

We will read the two remaining pieces of data. 

Again, albeit the step-by-step description, we encourage you to use method chaining.

<!-- grade_id: cell-030bd11a9b6c76c8, points: 0, grade: false, solution: true -->
```python
def make_works():
    # Read the works data.
    # works: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Remove the intervals.
    # works: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Select the columns GUID, ComposerName, WorkTitle, Movement and ConductorName.
    # YOUR CODE HERE
    raise NotImplementedError()
    return works


def make_soloists():
    # Read the soloists data and drop ProgramID, WorkID and MovementID.
    # YOUR CODE HERE
    raise NotImplementedError()
    return soloists


works = make_works()
soloists = make_soloists()
```

<!-- grade_id: cell-2a37eb128ea7e979, points: 2, grade: true, solution: false -->
```python
shape = str(works.shape)
expected_hash = 'cad58aa6cd33cfa24c08a0f0f846877178ab31278f212c80b16b952d9416f883'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(works.sample(random_state=0))
expected_hash = 'a27f5e63c23dcbe3db496bf7e1d3f7bb3ae9d3cd2cd1ff1d0498c73c66bd5703'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)

shape = str(soloists.shape)
expected_hash = 'a7b0d20a45ff1344e0398eebb162af9afb8805082b0dfdcb70e9a4b78f94dd13'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(soloists.sample(random_state=0))
expected_hash = '3352eb27df2a6dba982be1001af1af39a3adc031134f44fe14c90f34cdd1493e'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

## 5 Combine Works and Soloists (graded)

Like we did for Programs and Concerts, now we combine Works and Soloists.

<!-- grade_id: cell-3c2e4becb723a4a3, points: 0, grade: false, solution: true -->
```python
# Combine both dataframes, again using an inner type of join.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-13616b91f6c53cad, points: 2, grade: true, solution: false -->
```python
shape = str(works_and_soloists.shape)
expected_hash = 'c0e73877aac4f3916267cb58f2f122ffef32c79039bde2ecb217fda123270d12'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(works_and_soloists.sample(random_state=0))
expected_hash = 'e604d9c53131c6aabe798ef85e07413fbdd9090cc7e2d4fe911e6e23a0b5e7a8'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

## 6 Combine everything (graded)

The final goal here is to create a single data frame.

<!-- grade_id: cell-081458c0c0a40de2, points: 0, grade: false, solution: true -->
```python
# Combine everything into a single dataframe.
# nyp = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a0ca421b4504282f, points: 2, grade: true, solution: false -->
```python
shape = str(nyp.shape)
expected_hash = '9f690a900c6373cf2cb0ccf8748b0e63fdbf9ca201542794e3ef952cccdb4da3'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(nyp.sample(random_state=0))
expected_hash = '5e0a8c4e71b3a0ecb7c81b6719a2f9d9a4d1adc2ca34cd763a758ab6a1867c4c'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

## 7 Final transformations (graded)

Now, we perform the train-test split.

We also perform some final transformations on both datasets:
* Include some date features: Year, Month, Day and Weekday
* Drop Date, Season and GUID
* Change the column name Orchestra to OrchestraName, for consistency with other name columns
* Filter out composers that appear in less than 100 concerts.

<!-- grade_id: cell-f8ecf5eba7d6c38b, points: 0, grade: false, solution: true -->
```python
def preprocess_data(df):
    # You should follow these exact steps:
    #   1 - add_date_features, ideally using df.pipe
    #   2 - drop Date, Season and GUID
    #   3 - rename Orchestra to OrchestraName
    #   4 - filter out composers with less that 100 concerts (>= 100 rows)
    # YOUR CODE HERE
    raise NotImplementedError()
    return df

def add_date_features(df):
    # YOUR CODE HERE
    raise NotImplementedError()
    return df


nyp_ = preprocess_data(nyp)
X_train, X_test = train_test_split(nyp_, random_state=0)
```

<!-- grade_id: cell-77f754825659deb6, points: 2, grade: true, solution: false -->
```python
shape = str(nyp.shape)
expected_hash = '9f690a900c6373cf2cb0ccf8748b0e63fdbf9ca201542794e3ef952cccdb4da3'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(nyp.sample(random_state=0))
expected_hash = '5e0a8c4e71b3a0ecb7c81b6719a2f9d9a4d1adc2ca34cd763a758ab6a1867c4c'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)

shape = str(nyp_.shape)
expected_hash = '0d9fc79d1c283ef411b3d7a0d94d2c9077526ce2ca95526a55c3eb44fd321b3d'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(nyp_.sample(random_state=0))
expected_hash = '6fb60e43e3a7b600032305c93743963e6e2578cc8b1d6afd11b80fa9ac2daa57'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)

columns = str(nyp_.columns.values)
expected_hash = '7d131b98b4d7094443c094603c6db00aa20a79e49661acdefb33bf5fc1c071fa'
assert(hashlib.sha256(columns.encode()).hexdigest() == expected_hash)
```

And, finally, we would be ready to explore modeling.

For the next part, however, we will be using the famous [Boston House Prices Dataset](https://archive.ics.uci.edu/ml/machine-learning-databases/housing/housing.names).

## 8 Scaling features (graded)

About the Boston dataset:

> Each record in the database describes a Boston suburb or town. The data is from the Boston Standard Metropolitan Statistical Area (SMSA) in 1970.

The features are all numerical (real, positive):
* **CRIM** - per capita crime rate by town
* **ZN** - proportion of residential land zoned for lots over 25,000 sq.ft.
* **INDUS** - proportion of non-retail business acres per town
* **CHAS** - Charles River dummy variable (= 1 if tract bounds river; 0 otherwise)
* **NOX** - nitric oxides concentration (parts per 10 million)
* **RM** - average number of rooms per dwelling
* **AGE** - proportion of owner-occupied units built prior to 1940
* **DIS** - weighted distances to five Boston employment centres
* **RAD** - index of accessibility to radial highways
* **TAX** - full-value property-tax rate per \$10,000
* **PTRATIO** - pupil-teacher ratio by town
* **B** - 1000(Bk - 0.63)^2 where Bk is the proportion of blacks by town
* **LSTAT** - % lower status of the population
* **MEDV** - Median value of owner-occupied homes in \$1000's.

We want to scale all features to the same range, using `sklearn.preprocessing.MinMaxScaler()`.

<!-- grade_id: cell-520f1d8e99c15a49, points: 0, grade: false, solution: true -->
```python
boston = load_boston()
X = pd.DataFrame(data=boston.data, columns=boston.feature_names)
y = boston.target

X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=0)

# Initialize the MinMaxScaler to a [0, 5] range.
# YOUR CODE HERE
raise NotImplementedError()

# Fit on the training set and transform X_train. We expect X_train_
# to be a dataframe **just like** X_train, only scaled. 
# X_train_: pd.DataFrame = ...
# YOUR CODE HERE
raise NotImplementedError()

# Transform the test set.
# X_test_: pd.DataFrame = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-39a52095c35ce047, points: 2, grade: true, solution: false -->
```python
shape = str(X_train_.shape)
expected_hash = '6f696c7e30c15aae3f0fa4807b596cf15d28cadaf33602d8d20368f7ac921f26'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(X_train_.sample(random_state=0))
expected_hash = '158be9e6cfe3865a8ea5baffb8b8cc231030aa07bc114539de0df244e91ac5ec'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
_
columns = str(X_train_.columns.values)
expected_hash = 'c4e20218e7e33f0e771a608bb05ece0152f5a15fc6a0629b6c88cef7790fbfe1'
assert(hashlib.sha256(columns.encode()).hexdigest() == expected_hash)

shape = str(X_test_.shape)
expected_hash = 'aa2b4e3c1e358b4b9f21c2c86bbf1187020582395419f1a02a949d7a6efac9e4'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(X_test_.sample(random_state=0))
expected_hash = 'e2ea8d5b9b195cc081124efca8af59372b4004364afa3f78b5a50ea13149f9bb'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)

columns = str(X_test_.columns.values)
expected_hash = 'c4e20218e7e33f0e771a608bb05ece0152f5a15fc6a0629b6c88cef7790fbfe1'
assert(hashlib.sha256(columns.encode()).hexdigest() == expected_hash)
```

## 9 Build a ColumnSelector transformer (graded)

There's a simple transformer that can be useful, from times to times, when modeling.

What we want is to build a transformer that returns the columns we select beforehand. 

This transformer could be used to determine what features go into modeling.

<!-- grade_id: cell-3d8545164af32b84, points: 0, grade: false, solution: true -->
```python
class ColumnSelector(BaseEstimator, TransformerMixin):
    # Implement the __init__ method.
    # Our ColumnSelector must be able to receive a parameter columns.
    # The default value for columns must be set to 'all', so we can
    # initialize it without any explicit parameters.
    # YOUR CODE HERE
    raise NotImplementedError()
        
    # There's no need for a fit method in this case, it does nothing.
    # We should be able to call fit without any explicit parameters.
    # Meaning: we should be able to call ColumnSelector.fit().
    # YOUR CODE HERE
    raise NotImplementedError()

    # Transform should return all columns if the parameter columns we
    # passed upon initialization is equal to 'all'. If a column or a
    # list of columns are passed, only those should be returned.
    # YOUR CODE HERE
    raise NotImplementedError()
        

cols = ['CRIM', 'DIS', 'INDUS', 'RM', 'DIS', 'TAX', 'B']
selector = ColumnSelector(columns=cols)
X_train__ = selector.fit_transform(X_train_)
X_test__ = selector.transform(X_test_)
```

<!-- grade_id: cell-6d54a8e7ed69bd97, points: 2, grade: true, solution: false -->
```python
assert(ColumnSelector())
assert(selector.fit())

shape = str(X_train__.shape)
expected_hash = '5d4f688e84beb21ec07f136c16a6cc11318d4f5de7b81bf0232e5282d9834123'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(X_train__.sample(random_state=0))
expected_hash = '410bd230a9e60a209d5f2a9f705e41534149d7de98a0c0a4607d9c4121250437'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)

columns = str(X_train__.columns.values)
expected_hash = '901009bce1feeeccadd8cd499664598ff9319641e55dcda17a650c13c0626604'
assert(hashlib.sha256(columns.encode()).hexdigest() == expected_hash)

shape = str(X_test__.shape)
expected_hash = '0aba1c19151f76aa2ecb00fd75be05c6f73860573972e967f3d1fe1c44ae2629'
assert(hashlib.sha256(shape.encode()).hexdigest() == expected_hash)

sample = str(X_test__.sample(random_state=0))
expected_hash = '38a69d5837865bb4d5d6f6008331b4ae3449cee61aa10609b7db5fa73029c518'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)

columns = str(X_test__.columns.values)
expected_hash = '901009bce1feeeccadd8cd499664598ff9319641e55dcda17a650c13c0626604'
assert(hashlib.sha256(columns.encode()).hexdigest() == expected_hash)
```

## 10 Building the pipeline (graded)

Finally, we want to use the two transformers together and run a linear regression on top.

<!-- grade_id: cell-c6cf45c70a05aa5c, points: 0, grade: false, solution: true -->
```python
# Create a pipeline including:
#   1 - 'selector', ColumSelector(columns=cols)
#   2 - 'min_max', MinMaxScaler() with same range as above
#   3 - 'model', LinearRegression
# YOUR CODE HERE
raise NotImplementedError()


pipeline.fit(X_train, y_train)
y_pred = pipeline.predict(X_test)

mse = mean_squared_error(y_test, y_pred)
mae = mean_absolute_error(y_test, y_pred)
print('MSE: {}'.format(mse))
print('MAE: {}'.format(mae))
```

<!-- grade_id: cell-b520a77a388d20b2, points: 2, grade: true, solution: false -->
```python
assert(type(pipeline) == Pipeline)

params = str(pipeline.get_params)
expected_hash = '06f6c462bed73d9ff134dbfd1a8081c2c9f7e5e61e8166b5e14f4a3cf9af63ff'
assert(hashlib.sha256(params.encode()).hexdigest() == expected_hash)

sample = str(pd.Series(y_pred).sample(random_state=0))
expected_hash = 'f3ce987cb9aa2ef4caf2b86400515836c2ec8e2b620fbd94170f449f885f0195'
assert(hashlib.sha256(sample.encode()).hexdigest() == expected_hash)
```

Exercises complete, congratulations! You are about to become a certified data wrangler.